```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment1.ipynb")
```

## Introduction to Pandas and Scikit-Learn

Pandas is a powerful data manipulation and analysis library for Python. It provides data structures like DataFrames and Series that allow for efficient handling of structured data. Pandas is particularly useful for tasks such as reading and writing data in various formats, data cleaning, merging datasets, and performing complex operations on data.

Scikit-learn, on the other hand, is a machine learning library for Python. It provides a wide range of supervised and unsupervised learning algorithms, as well as tools for model selection, evaluation, and preprocessing. Scikit-learn is designed to be user-friendly and efficient, making it a popular choice for both beginners and experienced data scientists.

Together, Pandas and Scikit-learn form a powerful combination for data analysis and machine learning tasks. Pandas is often used to prepare and manipulate data, which can then be fed into Scikit-learn models for training and prediction.

In this assignment, we'll start with the fundamentals of data loading/manipulatin in pandas, then move on to basics of scikit-learn.

## Pandas

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

```python
df = pd.read_csv('adult.csv')
```

#### Let's take a quick look at the dataset. We'll use the .head() function to view the first 5 records of the dataset.

```python
df.head()
```

## Part 1: Data Loading and pre-processing

### Feel free to refer to the course notes on [pandas](https://tools4ds.github.io/DS701-Course-Notes/02B-Pandas.html) for the following questions.

**Question 1a**: Write a function `get_rows_and_columns` that takes as input a CSV filename, loads this file into a Pandas dataframe, and returns a tuple of the number of rows and columns in the dataframe.

```python
def get_rows_and_columns(file_path):
    ...
```

```python
grader.check("q1a")
```

**Question 1b**: Write a function `compute_missing_percentage` that converts '?' to NaN and returns the percentage of missing data (i.e., NaNs) for each column in the dataset.

```python
def compute_missing_percentage(df):
    ...
```

```python
grader.check("q1b")
```

**Question 1c**: Write a function `unique_ms` that returns the number of unique marital-statuses present in the dataset.

```python
def unique_ms(df):
    ...
```

```python
grader.check("q1c")
```

**Question 1d**: Write a function `get_categorical_columns` that identifies and returns a list of all the **categorical** columns in the dataset.

```python
def get_categorical_columns(df):
    ...
```

```python
grader.check("q1d")
```

## Part 2: Exploratory data analysis and visualizations

<!-- BEGIN QUESTION -->

**Question 2a**: Write a function `plot_categorical_distribution` to plot the distribution of the column 'relationship'.

```python
def plot_categorical_distribution(df):
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2b**: Write a function `plot_age_hours_scatter` that creates a scatter plot of 'age' vs 'hours-per-week', coloring points by 'income'.

```python
def plot_age_hours_scatter(df):
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2c**: Write a function `plot_income_by_education` which plots a stacked bar chart that shows the proportion of income levels for each 'education' category.

```python
def plot_income_by_education(df):
    ...
```

<!-- END QUESTION -->

## Part 3: Advanced Pandas Operations

**Question 3a**: Write a function `education_stats` that returns a dataframe with mean 'age' and median 'hours-per-week' categorized on the 'education' level.

```python
def education_stats(df):
    ...
```

```python
grader.check("q3a")
```

**Question 3b**: Write a function `calculate_high_income_percentage` that returns a dataframe of the percentage of individuals earning >50K for each 'native-country' and order them in descending order.

```python
def calculate_high_income_percentage(df):
    ...
```

```python
grader.check("q3b")
```

**Question 3c**: Write a function `occupations_by_hours` that returns a dataframe of the top 5 occupations with the higheset average 'hours-per-week'.

```python
def top_5_occupations_by_hours(df):
    ...
```

```python
grader.check("q3c")
```

## Scikit-learn

```python
# MAKE SURE TO RUN THIS CELL!

import warnings
from sklearn.preprocessing import MinMaxScaler
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_squared_error, mean_absolute_error

df_new = df.drop(columns=['native-country', 'fnlwgt']).dropna()
```

In this part, we'll implement [logistic regression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html) using scikit-learn. Logistic regression is used when trying to predict a binary outcome (0 or 1, True or False, etc.) We will go over the details of logistic regression in details later in the course. 

Here, we will try to predict income (>50k or <= 50k>) and follow standard ML procedures for data pre-processing. You can use scikit-learn's documentation, [the lecture notes on scikit-learn](https://tools4ds.github.io/DS701-Course-Notes/02C-Sklearn.html) or online resources for guidance. 

#### From here on use the 'df_new' variable instead of 'df'

We saw in lecture that models are trained on the 'training' set and evaluated on unseen data on the 'testing' set. The dataset has 'feature' (`X_train`, `X_test`) and the 'outcome' (`y_train`, `y_test`) variables. 

**Question 4a:** Write a function called `split_data` that takes a dataframe as its only parameter, splits it into training and test splits and returns them. Use 20% for the testing set. 

Use `train_test_split` to produce the splits. Provide a `random_state` of 42 for reproducibility.

`split_data` should return 4 things: X_train, X_test y_train and y_test. To do that, you need to pass in the X *and* the y (income column) to `train_test_split`.

```python
def split_data(df):
    ...
```

```python
grader.check("q4a")
```

**Question 4b**: Write a function `preprocess_data` that takes X_train, y_train, X_test, and y_test as input (the splits we created earlier!) and does the following:

 - Scale the *numerical* columns using sklearn's `MinMaxScaler` to the range [0,1] for both train and test sets
  
 - Replace "<=50K" with 0 and ">50K" with 1 in both y_train and y_test

 - One-hot encode the categorical columns for both train and test sets. Check the next cell for some hints! 
  
The function should then return the preprocessed X_train, y_train, X_test, and y_test

Refer to the material below and [sklearn course notes](https://tools4ds.github.io/DS701-Course-Notes/02C-Sklearn.html#prepare-the-dataset) for help! 

One hot encoding is a way of turning textual data into numbers, so that models can work with them. 

Pandas has a method called `pd.get_dummies()` that can do one-hot encoding. Here's an example usage: 
  
   ```python
    # Sample DataFrame with categorical columns
    data = {'City': ['New York', 'Los Angeles', 'New York', 'Chicago'],
        'Gender': ['Female', 'Male', 'Male', 'Female']}

    df = pd.DataFrame(data) 

    df_encoded = pd.get_dummies(df, columns=['City', 'Gender']) # notice how I'm passing in the columns -- you should do this too! Hint: you wrote a function for this earlier. 
   ```
   
   And then `df_encoded` will look like: 

   ```python
      City_Chicago  City_Los Angeles  City_New York  Gender_Female  Gender_Male
    0             0                 0              1              1            0
    1             0                 1              0              0            1
    2             0                 0              1              0            1
    3             1                 0              0              1            0
   ```

   Notice how now every `City` value has its own column, and that every row with a city has a 1 for that city (row 1 in the old dataframe has New York for the `City`, and row in the new dataframe has a 1 for `City_New York`). Everywhere else you have a 0. 

```python
def preprocess_data(X_train, X_test, y_train, y_test):
    ...
```

```python
grader.check("q4b")
```

Let's get to training! Remember, we're trying to predict whether income is more then 50k (>50k), or less than or equal to (<=50k).

**Question 4c:** Write a function called train_model that takes the training splits (X_train and y_train) as its parameters. 

- Initialize the logistic regression model
- Fit it to our data. (Training step)

At the end, return the fitted model. 

You can refer to [LogisticRegression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html)

```python
def train_model(X_train, y_train):
    ...
```

```python
grader.check("q4c")
```

Let's evaluate the performance of our model! 

**Question 4d:** Write a function called evaluate_model that takes the fitted model and `X_test`, `y_test` as parameters, runs the model on the testing features (`X_test`) and returns the *accuracy score* of the predictions against the ground truth (`y_test`). 

You can refer to [sklearn metrics](https://scikit-learn.org/stable/modules/model_evaluation.html#regression-metrics).

```python
from sklearn.metrics import accuracy_score

def evaluate_model(model, X_test, y_test):
    ...
```

```python
grader.check("q4d")
```

## Submission

Submit the notebook after completing all the questions to Gradescope to view your results!

Tests for q1a (points: None)
```
>>> get_rows_and_columns('adult.csv') == (48842, 15)
True
```
```
>>> get_rows_and_columns('adult.csv') == (4, 15)
False
```


Tests for q1b (points: None)
```
>>> bool(np.isclose(compute_missing_percentage(df)['age'], 0.0))
True
```
```
>>> bool(np.isclose(compute_missing_percentage(df)['workclass'], 5.730724))
True
```


Tests for q1c (points: None)
```
>>> unique_ms(df) == 7
True
```
```
>>> unique_ms(df) == 6
False
```


Tests for q1d (points: None)
```
>>> get_categorical_columns(df) == ['workclass', 'education', 'marital-status', 'occupation', 'relationship', 'race', 'gender', 'native-country', 'income']
True
```
```
>>> get_categorical_columns(df) == ['workclass', 'education', 'marital-status', 'occupation', 'relationship', 'race', 'gender', 'native-country']
False
```


Tests for q3a (points: None)
```
>>> len(education_stats(df)) == 16
True
```
```
>>> set(education_stats(df).columns) == {'education', 'age', 'hours-per-week'}
True
```
```
>>> bool(np.isclose(education_stats(df).iloc[0]['age'], 37.902808))
True
```
```
>>> bool(education_stats(df).iloc[1]['hours-per-week'] == 40.0)
True
```
```
>>> int(education_stats(df)['hours-per-week'].value_counts().get(40.0, 0)) == 14
True
```
```
>>> education_stats(df).iloc[13]['education'] == 'Preschool'
True
```
```
>>> education_stats(df).iloc[7]['education'] == 'Assoc-acdm'
True
```


Tests for q3b (points: None)
```
>>> bool(np.isclose(calculate_high_income_percentage(df)['France'], 42.10526315789473))
True
```
```
>>> bool(np.isclose(calculate_high_income_percentage(df)['India'], 41.059603))
True
```
```
>>> bool(np.isclose(calculate_high_income_percentage(df)['China'], 29.508197))
True
```


Tests for q3c (points: None)
```
>>> bool(np.isclose(top_5_occupations_by_hours(df)['Farming-fishing'], 46.81744966442953))
True
```
```
>>> bool(np.isclose(top_5_occupations_by_hours(df)['Exec-managerial'], 44.975353))
True
```
```
>>> bool(np.isclose(top_5_occupations_by_hours(df)['Prof-specialty'], 42))
False
```


Tests for q4a (points: None)
```
>>> def test_split_data():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     output = split_data(df_new)
...     assert len(output) == 4, 'split_data() should return 4 things. Make sure you pass both the X and the y to train_test_split.'
...     X_train, X_test, y_train, y_test = output
...     assert len(X_train) + len(X_test) == len(df_new), 'The splits do not add up to the original data size.'
...     assert len(y_train) + len(y_test) == len(df_new), 'The splits do not add up to the original data size.'
...     assert np.all(X_train.columns == X_test.columns), 'The columns in X_train and X_test are not the same.'
...     assert len(X_train) > 0 and len(X_test) > 0, 'Some of the splits are empty. Make sure you pass the correct values to train_test_split'
...     assert len(y_train) > 0 and len(y_test) > 0, 'Some of the splits are empty. Make sure you pass the correct values to train_test_split'
>>> test_split_data()
```


Tests for q4b (points: None)
```
>>> def test_correct_X():
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         X_train, X_test, y_train, y_test = split_data(df_new)
...         X_train, X_test, y_train, y_test = preprocess_data(X_train, X_test, y_train, y_test)
...     assert np.all(X_train.columns == X_test.columns), 'The columns in X_train and X_test are not the same.'
...     assert np.allclose(np.mean(X_train), np.mean(X_test), atol=0.01), 'The mean of every column is not the same in X_train and X_test.'
...     assert np.all((X_train >= 0) & (X_train <= 1)), 'X_train has values that are not between 0 and 1.'
...     assert np.all((X_test >= 0) & (X_test <= 1)), 'X_test has values that are not between 0 and 1.'
...     assert np.all((np.mean(X_train) >= 0) & (np.mean(X_train) <= 1)), 'X_train has columns with mean values not between 0 and 1.'
...     assert np.all((np.mean(X_test) >= 0) & (np.mean(X_test) <= 1)), 'X_test has columns with mean values not between 0 and 1.'
>>> 
>>> def test_correct_y():
...     splits = train_test_split(df_new.drop(columns='income'), df_new['income'], test_size=0.2, random_state=42)
...     X_train, X_test, y_train, y_test = splits
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         X_train, X_test, y_train, y_test = preprocess_data(X_train, X_test, y_train, y_test)
...     assert np.all(y_train.isin([0, 1])), 'y_train has values that are not 0 or 1.'
...     assert np.all(y_test.isin([0, 1])), 'y_test has values that are not 0 or 1.'
>>> test_correct_X()
>>> test_correct_y()
```


Tests for q4c (points: None)
```
>>> def test_train_model():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     X_train, X_test, y_train, y_test = split_data(df_new)
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         X_train, X_test, y_train, y_test = preprocess_data(X_train, X_test, y_train, y_test)
...     model = train_model(X_train, y_train)
...     return isinstance(model, LogisticRegression)
>>> test_train_model()
True
```


Tests for q4d (points: None)
```
>>> def test_evaluate_model():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     X_train, X_test, y_train, y_test = split_data(df_new)
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         X_train, X_test, y_train, y_test = preprocess_data(X_train, X_test, y_train, y_test)
...     model = train_model(X_train, y_train)
...     model = train_model(X_train, y_train)
...     accuracy = evaluate_model(model, X_test, y_test)
...     return bool(np.isclose(accuracy, 0.85, atol=0.1))
>>> test_evaluate_model()
True
```
